(assume |Goal| (not (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_))))))
(step t2 (cl (= (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)))) true)) :rule or_simplify)
(step t3 (cl (= (not (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_))))) (not true))) :rule cong :premises (t2))
(step t4 (cl (= (not true) false)) :rule not_simplify)
(step t5 (cl (= (not (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_))))) false)) :rule trans :premises (t3 t4))
(step t6 (cl (not (= (not (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_))))) false)) (not (not (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)))))) false) :rule equiv_pos2)
(step t7 (cl (not (not (not (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_))))))) (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_))))) :rule not_not)
(step t8 (cl (not (= (not (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_))))) false)) (or (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)) (not (smt__TLA__Mem smt__CONSTANT_x_ (smt__TLA__SetSt_flatnd_1 smt__CONSTANT_S_)))) false) :rule th_resolution :premises (t7 t6))
(step t9 (cl false) :rule th_resolution :premises (|Goal| t5 t8))
(step t10 (cl (not false)) :rule false)
(step t11 (cl) :rule resolution :premises (t9 t10))
